# <span style="color:teal;">CIS 211 Project 8:  &nbsp; Bears and Fish</span>

##### Due 11:00 P.M. Thursday June 8

##### <span style="color:red">Group Members:</span>

For this project you are allowed to work in small groups of up to 3 people.

If you work in a group only one group member should upload this notebook to Canvas.  Edit this cell to include the name and DuckID for each member of the group:

**Name:**

**Name:**

**Name:**

##  <span style="color:teal;">Overview</span> 

The last project this term is a cellular automaton style biological simulation.  

The sytem is described in Chapter 11 of the Miller and Ranum text.  **Read the description of the simulation in the textbook before you start working on this project.**  Your code will be quite a bit different than the code in the book, but the rules of the simulation are the same, and you need to know how objects interact.

The project has three main parts:
* A class named World that defines a 2D grid where cells are accessed according to their row and column coordinates
* Classes for Bear and Fish, the organisms that can inhabit the world
* Two top level functions, `wbf` and `step_system`; the first will make a world and populate it at random with bears and fish, and the second will run a simulation using that world.

We're written `step_system` for you -- it's in the last code cell in this notebook -- but you need to design, implement, and test the classes and the `wbf` function.

##  <span style="color:teal;">Libraries</span> 

Your program will need to use `numpy` (a numeric processing library) and `random` (the builtin random number generator module).  You can import additional modules if you wish; add the import statements to this code cell.

```python
IPython = (__doc__ is not None) and ('IPython' in __doc__)
Main    = __name__ == '__main__'

import numpy as np
import random
```

##  <span style="color:teal;">Event Log</span> 

A useful debugging technique is to save descriptions of events in a log.  Any code in the simulator, whether it is the top level function or a method in one of your classes, can call a function named `log`, passing it a description of an event that just occurred.  Examples might be 
```
new Fish in cell (x,y)
```
or 
```
Bear in cell (x,y) eats Fish in cell (x,y)
```

The function below uses a global variable named `logging`.  If you want to take advantage of the `log` function:
* edit the cell to set `logging` to True
* add calls to `log` at various points in your methods

Later, when doing large scale simulations, you can turn off logging simply by setting `logging` to False before you call the top level simulation function.

**Note:** &nbsp; you can change the definition of the `log` function however you want, _e.g._ you can have it write log messages to a file or print additional information.

```python
logging = False

def log(message):
    if logging:
        print(message)
```

##  <span style="color:teal;">World (20 points)</span>

Write the code for the World class in the code cell below.

The constructor will be passed the grid size (number of rows and columns), and it should initialize all the cells to `None`.

Define the following methods:
* `add` will be passed an object and a location; it should store the object in the specified location 
* `fetch` should return the current contents of a location (which may be `None`)
* `remove` should set a location to `None`.

Finally, define a method named `biota` that will return a list of all items currently in the world.

You can have additional instance methods, class variables, or class methods.  Make sure you describe any new additions in the documentation.

##### <span style=color:red>Code</span> 

**Important** Write the definition of your World class in the following code cell. Do not delete or move this cell.

<!-- grade_id: world, points: 10, grade: true, solution: true -->
```python
class World:
    ''' The world class defines the dimensions of the grid in which our
    other bear and fish objects are created in. World objects have the 
    methods add, fetch, remove, and biota. The constructor is passed 
    a grid size and intializes all the cells within the grid to None.
    '''
    def __init__(self, x, y):
        ''' The constructor function is passed x and y dimensions
        to create the grid.
        '''
        self._x = x
        self._y = y
        self._grid = np.array([None] * (x*y)).reshape(x,y)
        
    def x(self):
        ''' This is a getter method for the instance variable x'''
        
        return self._x
    
    def y(self):
        ''' This is a getter method for the instance variable y'''
        
        return self._y
   
    def __getitem__(self, loc):
        ''' This method gives world objects the indexing operation
        so the we can access content at a given location within our
        grid
        '''
        
        return self._grid[loc]
    
    def __setitem__(self, loc, val):
        ''' This method allows us to assign a value to a cell at a 
        given location within our grid.
        '''
        
        self._grid[loc] = val
        
    def add(self, obj, loc):
        ''' The add method is passed an object and a location for the
        object passed to be placed at on the grid constructed.
        '''
        
        self._grid[loc] = obj
        
    def fetch(self, loc):
        ''' The fetch method returns the contents of a cell at the given
        location on the grid. If nothing is in the cell, None should
        be returned.
        '''
        
        return self._grid[loc]
    
    def remove(self, obj, loc):
        ''' The remove method is passed a location and the cell in the
        grid at the passed location is set to None.
        '''
        
        self._grid[loc] = None
        
    def biota(self):
        ''' The method biota returns the contents of all the cells
        within the grid
        '''
        log(self._grid)
        life = []
        for obj in self._grid.flat:
            if obj != None:
                life.append(obj)
        
        return life
```

##### <span style="color:red">Tests:</span>

Use the following code cell as a "sandbox" if you want to do your own tests.  You can add additional cells here if you want.

```python

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: world_test_1, points: 1, grade: true, solution: false -->
```python
# A new world has no objects
w0 = World(5,5)

assert len(w0.biota()) == 0
```

<!-- grade_id: world_test_2, points: 2, grade: true, solution: false -->
```python
# Test the add and fetch operators
w1 = World(5,5)

w1.add('hello', (0,0))
w1.add('world', (1,1))

assert w1.fetch((0,0)) == 'hello'
assert w1.fetch((1,1)) == 'world'
```

<!-- grade_id: world_test_3, points: 2, grade: true, solution: false -->
```python
# Test the remove and biota methods
w2 = World(5,5)

w2.add('hello', (0,0))
w2.add('world', (1,1))

assert sorted(w2.biota()) == ['hello','world']

w2.remove('hello', (0,0))
w2.remove('world', (1,1))

assert w2.biota() == []
```

##### <span style=color:red>Documentation</span> 

**Important** Write your documentation in the following markdown cell.  Do not delete or move this cell.

<!-- grade_id: world_doc, points: 5, grade: true, solution: true -->
The World class is initialized with the instance variables x,y, and grid. The x and y are passed to the grid to make an array that's saved in our grid instance variable so that it can be referenced by
methods later. 
setitem and getitem were also added to our World object so that it 
could be indexed and items in the object could be set to other values.
The rest of the methods were defined and documented above.
add, fetch, remove, and biota.
Along with getter methods for our x and y instance variables.

## <span style="color:teal;">Fish and Bears</span>

To run a simulation we need to add a random collection of animals to the world.  The two types of animals in this simulation are fish and bears, and you will write class definitions named Fish and Bear that implement the behaviors of the animals.

### Fish Class

During the simulation a Fish object needs to behave as follows:

(1) Fish are susceptible to overcrowding:  if there are fish in 2 or more neighboring cells the fish dies (it's removed from the simulation)

(2) A fish can reproduce if it has been alive for a certain number of time steps: a random neighboring cell is chosen, and if that cell is empty, a new fish is placed in that cell

(3) A fish can move to another cell:  it picks a random direction, and if the neighboring cell in that direction is unoccupied the fish moves there

The constructor for the Fish class will be passed a reference to a World object and a location, in the form of a tuple with a row number and column number (the object needs to know its location so it can look for other objects in neighboring cells).

The class should include the following methods:
* `live` implements rules 1 and 2 shown above
* `move` implements rule 3
* `location` returns the current grid location (row and column) of the object

Define a class variable named `breed_interval` to specify how many time steps a fish must be alive before it reproduces; the initial value for this variable is 12.


### Bear Class

During the simulation a Bear object needs to behave as follows:

(1) A bear looks for fish in each adjacent cell; if it finds one or more fish it eats one at random 

(2) If a bear has not eaten for certain number of time steps it dies (it's removed from the simulation)

(3) A bear can reproduce if it has been alive for a certain number of time steps: a random neighboring cell is chosen, and if that cell is empty, a new bear is placed in that cell

(4) A bear can move to another cell:  it picks a random direction, and if the neighboring cell in that direction is unoccupied the bear moves there

Define a class variable named `breed_interval` to specify how many time steps a bear must be alive before it reproduces; the initial value for this variable is 8.  Define another class variable named `survive_without_food` to be the number of time steps a bear can live before it dies from starvation; the initial value for this variable is 10.

The constructor will be passed a reference to a World object and a location, in the form of a tuple with a row number and column number (the object needs to know its location so it can look for other objects in neighboring cells).

The class should include the following methods:
* `live` implements rules 1, 2, and 3 shown above
* `move` implements rule 4
* `location` returns the current grid location (row and column) of the object

### Animal (Base Class)?

From the descriptions above it should be apparent that fish and bears have some things in common.

For **full credit** on the coding and documnetation portions of this project you should define a class named Animal and use it as the base class for your Fish and Bear classes.  Some things to think about as you design your classes:
* are there behaviors or operations that are common to both and that can be implemented just once in Animal?
* perhaps a behavior or operation can be defined with a default in the Animal class, and then overridden in the derived class?

You can still receive **partial credit** if you skip the Animal class and simply write completely separate Fish and Bear classes.  None of the unit tests assume there is a class named Animal.

### Details and Hints

**Incremental Development:**  One strategy you might consider is to write a complete implementation for one class, either Bear or Fish.  After you have debugged the class and it passes its unit tests you'll have a better idea of what to implement in the other class.  Then you can start moving common behaviors to the Animal class while you are writing the other derived class.

**Fish:** &nbsp; The way the simulation is defined a fish might be eaten before the top level simulation calls the `live` method.  Your `live` method should check to make sure the fish is still alive.  The easiest way to do this is to include an instance variable named `_alive` that is set to True when the fish is initialized and set to False when it dies.

You can have additional instance methods, class variables, or class methods.  Make sure you describe any new additions in the documentation.

**Animal:** &nbsp; Keep this in mind when you design the class hierarchy: every object has an attribute named ``__class__`` (with two underscores before and after the name).  It is a reference to the class an object was defined with.  As an example of how to use it, consider what would happen if we want to define `reproduce` in the Animal class so it is inherited by both Fish and Bear classes.  We need to know the value of `breed_interval` in each object's own class.  We can find this value using the expression
```
self.__class__.breed_interval
```
This will be a reference to `Fish.breed_interval` or `Bear.breed_interval`, depending on whether a `reproduce` was called with Fish object or Bear object.


### <span style="color:teal;">Animal (15 points)</span>

##### <span style=color:red>Code</span> 

If you implement the Animal class write the definition in the following code cell

**Important:** &nbsp; Do not delete or move this cell.

<!-- grade_id: animal, points: 10, grade: true, solution: true -->
```python
class Animal:
    ''' The Animal class is a object that can either be inheritted by the
    Bear or Fish class. This class encompasses all the shared instance
    variables and methods between our Bear and Fish classes.
    '''
    
    def __init__(self, world, location):
        ''' This constructor intializes the object to be living, the
        breeding interval of an animal, the location of that animal
        on the grid world, and the time the animal has been alive.
        '''
        
        self._grid = world
        self._alive = True
        self._loc = location
        self._time_alive = 0
        self._grid.add(self, location)
        
    def reproduce(self):
        ''' This function checks the amount of time steps the object
        has been alive and compares the value to the objects breed_interval.
        If the amount of time steps is greater then the object chooses a
        random cell using a callback function, if the cell is empty the 
        animal reproduces.
        '''
        log('reproducing')
        loc = self.random_loc()
        if (self.__class__.breed_interval <= self._time_alive) and (self._grid.fetch(loc) == 'None'):
            self.__class__(self._grid, loc)
        
        
    def move(self):
        ''' This method checks a randomly picked neighboring cell, using
        a reference to the method random_loc, of the object and determines
        if the checked cell is occupied, if not the object moves to that 
        location, otherwise the object does not moce.
        '''
        loc = self.random_loc()
        if self._grid.fetch(loc) == 'None':
            self._loc = loc
        
        
    def location(self):
        ''' This method returns the grid location as a tuple (row and column)
        of the object
        '''
        
        return self._loc
    
    
    def random_loc(self):
        ''' This method generates a random location and returns the value
        so that it can be used as a callback for other methods
        '''
        x = random.randint(self._loc[0]-1, self._loc[0]+1)
        y = random.randint(self._loc[1]-1, self._loc[1]+1)
        
        return (x,y)
    
    def f_neighbors(self):
        ''' This method counts the number of Fish objects neighboring'''
        
        
        locations = [(self._loc[0]-1,self._loc[1]), (self._loc[0]-1,self._loc[1]+1), (self._loc[0]-1,self._loc[1]-1), (self._loc[0],self._loc[1]+1), (self._loc[0],self._loc[1]-1), (self._loc[0]+1,self._loc[1]), (self._loc[0]+1,self._loc[1]+1), (self._loc[0]+1,self._loc[1]-1)]
        fish = 0
        for loc in locations:
            log(self._grid.fetch(loc))
            if self._grid.fetch(loc) == Fish.__class__:
                log('Fish at {}'.format(loc))
                fish = fish + 1
        log(fish)
        return fish
                          
                          
                          
```

##### <span style="color:red">Tests:</span>

There are **no unit tests for the Animal class.**  If you implement this class it will be tested automatically as part of the unit tests for Bear and Fish.

You can still make your own tests of the base class if you wish.  Add expressions that create and test Animal objects to the following "sandbox" cells.

```python

```

##### <span style=color:red>Documentation</span> 

If you implement the Animal class write your documentation in the following markdown cell.  

**Important** Do not delete or move this cell.

<!-- grade_id: animal_doc, points: 5, grade: true, solution: true -->
The Animal class encompasses all the shared attributes and methods of the Bear and Fish classes. Such as whether or not the animal is alive, the grid world the animal is placed within, the time the animal has been alive, and the location within the grid.
The methods include location, move, reproduce, and two of my own methods f_neighbors and random_loc. f_neighbors calculates the number of fish objects near the objects location and random_loc picks a random location adjacent to the object and returns the tuple of coordinates. 

### <span style="color:teal;">Fish (20 points)</span>

##### <span style=color:red>Code</span> 

**Important:** &nbsp; Write the Python code for your Fish class in the following code cell.  Do not delete or move this cell.

<!-- grade_id: fish, points: 7, grade: true, solution: true -->
```python
class Fish(Animal):
    ''' The Fish class constructs an object that is placed randomly on
    a grid that's a World object. The Fish object starts out alive and 
    can move, breed, and of course die by overcrowding or by being eaten
    by a Bear object.
    '''
    
    breed_interval = 12
    
    def __init__(self, world, location):
        ''' This method inherits from the Animal class and utilizes the
        the Animal constructor and initializes the breed interval specific
        to the Fish object.
        '''
        
        Animal.__init__(self, world, location)
        self._label = 'Fish'
        
    def __repr__(self):
        
        return self._label
    
    def live(self):
        ''' This method checks for overcrowding of other fish and 
        also references the Animal class method reproduce.
        '''
        self._time_alive += 1
        
        if self.f_neighbors() > 1:
            self._alive = False
            self._grid.remove(self._grid.fetch(self._loc), self._loc)
            return
        
        self.reproduce()
        
        
        
```

##### <span style="color:red">Tests:</span>

Use the following code cell as a "sandbox" if you want to do your own tests.  You can add additional cells here if you want.

```python

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: fish_test_1, points: 2, grade: true, solution: false -->
```python
# A new world has no objects
fw1 = World(5,5)
assert len(fw1.biota()) == 0

# After adding a fish there should be one object
f1 = Fish(fw1, (2,2))
assert len(fw1.biota()) == 1

# Test the location method
assert f1.location() == (2,2)
```

<!-- grade_id: fish_test_2, points: 2, grade: true, solution: false -->
```python
# Setting breed_interval to 1 should cause a fish to reproduce when live is called
fw2 = World(5,5)
Fish.breed_interval = 1
f2 = Fish(fw2, (2,2))
f2.live()
assert len(fw2.biota()) == 2

# Reset the interval to original value for remaining tests
Fish.breed_interval = 12
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-10-7cce60d68f23> in <module>()
      4 f2 = Fish(fw2, (2,2))
      5 f2.live()
----> 6 assert len(fw2.biota()) == 2
      7 
      8 # Reset the interval to original value for remaining tests

AssertionError:
```

<!-- grade_id: fish_test_3, points: 2, grade: true, solution: false -->
```python
# Make three fish, the one in the middle should die from overcrowding
fw3 = World(5,5)
f3 = Fish(fw3, (2,2))
Fish(fw3, (1,1))
Fish(fw3, (3,3))
f3.live()
assert len(fw3.biota()) == 2
```

<!-- grade_id: fish_test_4, points: 2, grade: true, solution: false -->
```python
# When a fish moves it should be within one cell of its original location
fw4 = World(5,5)
f4 = Fish(fw4, (2,2))
f4.move()
r, c = f4.location()
assert (r,c) != (2,2)
assert abs(r-2) <= 1 and abs(c-2) <= 1
```

##### <span style=color:red>Documentation</span> 

**Important:** &nbsp; Write the documentation for your Fish class in the following markdown cell.  Do not delete or move this cell.

<!-- grade_id: fish_doc, points: 5, grade: true, solution: true -->
The Fish class inherits all the variables and methods from the previously created Animal class. I added a overload to the print statement and gave the fish object a label within the grid. The fish also has a class variable breed_interval that's unique to this class.
The live method uses callback functions from our Animal class, such as reproduce and f_neighbors.

###  <span style="color:teal;">Bear (20 points)</span>

##### <span style=color:red>Code</span> 

**Important:** &nbsp; Write the Python code for your Bear class in the following code cell.  Do not delete or move this cell.

<!-- grade_id: bear, points: 7, grade: true, solution: true -->
```python
class Bear(Animal):
    ''' The Fish class constructs an object that is placed randomly on
    a grid that's a World object. The Fish object starts out alive and 
    can move, breed, and of course die by overcrowding or by being eaten
    by a Bear object.
    '''
    
    breed_interval = 8
    survive_without_food = 10
    
    def __init__(self, world, location):
        ''' This method inherits from the Animal class and utilizes the
        the Animal constructor and initializes the breed interval specific
        to the Fish object.
        '''
        
        Animal.__init__(self, world, location)
        self._label = 'Bear'
    
        
    def __repr__(self):
        
        return self._label
    
    def live(self):
        ''' This method checks for overcrowding of other fish and 
        also references the Animal class method reproduce.
        '''
        self._time_alive += 1
        
        if self.f_neighbors() > 1:
            locations = [(self._loc[0]-1,self._loc[1]), (self._loc[0]-1,self._loc[1]+1), (self._loc[0]-1,self._loc[1]-1), (self._loc[0],self._loc[1]+1), (self._loc[0],self._loc[1]-1), (self._loc[0]+1,self._loc[1]), (self._loc[0]+1,self._loc[1]+1), (self._loc[0]+1,self._loc[1]-1)]
            f_locs = []
            for loc in locations:
                if self._grid.fetch(loc) == Fish.__class__:
                    f_locs.append(loc)
                    
            choice = random.choice(f_locs)
            self._grid.remove(self._grid.fetch(choice), choice)
            self.survive_without_food = 10
        
        self.survive_without_food -= 1
        self.reproduce()
        
        
```

##### <span style="color:red">Tests:</span>

Use the following code cell as a "sandbox" if you want to do your own tests.  You can add additional cells here if you want.

```python

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: bear_test_1, points: 2, grade: true, solution: false -->
```python
# Test the Bear constructor and location method
bw1 = World(5,5)
b1 = Bear(bw1, (1,1))
assert len(bw1.biota()) == 1
assert b1.location() == (1,1)
```

<!-- grade_id: bear_test_2, points: 2, grade: true, solution: false -->
```python
# Repeat the reproduction test for Bears
bw2 = World(5,5)
Bear.breed_interval = 1
b2 = Bear(bw2, (2,2))
b2.live()
assert len(bw2.biota()) == 2
Bear.breed_interval = 8
```

<!-- grade_id: bear_test_3, points: 2, grade: true, solution: false -->
```python
# Make fish for the bear to eat, count the number of objects after eating
bw3 = World(5,5)
b3 = Bear(bw3, (2,2))
Fish(bw3, (1,1))
Fish(bw3, (3,3))
b3.live()
assert len(bw3.biota()) == 2
```

<!-- grade_id: bear_test_4, points: 2, grade: true, solution: false -->
```python
# Setting the survival limit to 1 should cause a Bear to starve 
bw4 = World(5,5)
Bear.survive_without_food = 1
b4 = Bear(bw4, (2,2))
b4.live()
assert len(bw4.biota()) == 0
Bear.survive_without_food = 10
```

<!-- grade_id: bear_test_5, points: 2, grade: true, solution: false -->
```python
# Repeat the move test for bears
bw5 = World(5,5)
b5 = Bear(bw5, (2,2))
b5.move()
r, c = b5.location()
assert (r,c) != (2,2)
assert abs(r-2) <= 1 and abs(c-2) <= 1
```

##### <span style=color:red>Documentation</span> 

**Important:** &nbsp; Write the documentation for your Bear class in the following markdown cell.  Do not delete or move this cell.

<!-- grade_id: bear_doc, points: 3, grade: true, solution: true -->
The Bear class inherits all the variables and methods from the previously created Animal class. I added a overload to the print statement and gave the bear object a label within the grid. The bear also has a class variable breed_interval and survive_withoout_food that's unique to this class.
The live method uses callback functions from our Animal class, such as reproduce and f_neighbors.

##  <span style="color:teal;">The `wbf` Function (15 points)</span>

Fill in the body of the `wbf` function so it returns a new World object with the specified number of rows and columns and with the specified number of Bear and Fish objects at random locations.

When we grade your project we will call `wbf` to make a World object and then use the main loop (implemented by `step_system`) to run the simulation.

<!-- grade_id: wbf, points: 9, grade: true, solution: true -->
```python
def wbf(nrows, ncols, nbears, nfish):
    ''' This function creates and saves our world in a variable grid
    then uses while loops to ensure that the bear and fish are placed in 
    empty cells throughout the grid at random locations.
    '''
    
    grid = World(nrows,ncols)
    n = nbears
    while n >= nbears:
        x = random.randint(0, nrows)
        y = random.randint(0, ncols)
        if grid.fetch(x,y) == None:
            grid.add(Bear(grid,(x,y)),(x,y))
            n -= 1
            
    m = nfish
    while m >= nfish:
        x = random.randint(0, nrows)
        y = random.randint(0, ncols)
        if grid.fetch(x,y) == None:
            grid.add(Fish(grid,(x,y)),(x,y))
            n -= 1
            
        
```

##### <span style="color:red">Tests:</span>

Use the following code cell as a "sandbox" if you want to do your own tests.  You can add additional cells here if you want.

```python

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: wbf_test, points: 3, grade: true, solution: false -->
```python
w = wbf(10,10,3,12)

dct = { Bear: 0, Fish: 0 }
for x in w.biota():
    dct[x.__class__] += 1

assert dct[Bear] == 3
assert dct[Fish] == 12
```

##### <span style=color:red>Documentation</span> 

**Important:** &nbsp; Write the documentation for your `wbf` function in the following markdown cell.  Do not delete or move this cell.

<!-- grade_id: wbf_doc, points: 3, grade: true, solution: true -->
I used the World class to intialize the grid specific to the size given
nrows by ncolumns. Then I used while loops with counters to add the bear and fish to the world at random locations. The loops also ensure that the bear and fish do not end up in overlapping cells.

##  <span style="color:teal;">The `step_system` Function</span>

We've written this function for you.  It will run a single time step of the simulation.  Pass it a World object containing a grid populated with Bear and Fish objects and it will (a) see which animals survive, then (b) allow all the animals to move to a new location.

```python
def step_system(world):
    for x in world.biota():
        x.live()
    for x in world.biota():
        x.move()
```

The code cell below is an example of how to run a complete simulation.  Call `wbf` to make a World populated by Bear and Fish objects, and then run the simulation for some number of iterations.

You can run this cell to do an "integration test" of your code.  You can also copy it and modify it for the exercises in the next section.

```python
if IPython:
    logging = True
    w = wbf(10,10,3,12)
    for i in range(3):
        print(w)
        step_system(w)
```

##  <span style="color:teal;">Experiments (10 points)</span>

Run some experiments with the top level simulation loop and describe the results in the markdown cell below.  Some ideas of things to try:
* The settings for the Bear class `breed_interval` and `survive_without_food` variables come from the textbook.  Will the world ever run out of bears with these settings?
* Change the settings so `Bear.breed_interval` is larger than `Bear.survive_without_food`.  How does that change the outcome?
* Set the `breed_interval` counter for the Fish class to a smaller number (e.g. 4) so the world has more fish.  What effect does that have?
* Write a function that runs the simulation for a specified number of generations, or until there are no more objects left in the grid.  What combination of parameters leads to the largest number of time steps before the simulation halts?

<!-- grade_id: experiments, points: 10, grade: true, solution: true -->
YOUR DOCUMENTATION HERE

##  <span style="color:teal;">Extra Credit Ideas</span>

Here are some suggestions for ways to extend the simulation.  We will consider other types of extensions -- send a request to `conery@uoregon.edu` with your proposal.

* Implement the Plant class described in Section 10.7 of the textbook, and modify the Fish class so fish eat plants and die if they don't find enough food.

* If you implement Plant, how does that affect the class hierarchy?  Is there an even more general class called Organism, with Plants and Animals subtypes of organism?

* Experiment with data structures: make a second version of the World class, but use a list-of-lists approach to making the grid.  Which is more efficient, a list of lists or a numpy array?  Which scales better when larger worlds are used in the simulation?

* Implement a GUI using `tkinter` that is similar to the Solar System GUI.  A canvas should display the world along with images for the Bear and Fish objects (you can download `Bear.gif` and `Fish.gif` from Canvas).  Use a spinbox or text entry box to specify the number of Bear and Fish objects and the number of time steps to run.  Include a Run button to start the simulation.


```python

```