# <span style="color:teal;">CIS 211 Notebook Demo</span>

** This notebook is for an in-class demonstration of how to use Jupyter for weekly programming projects. **

The programming exercises here are based on Project 2 (decrypting strings made with a substitution cipher).

##### Reading:  M&R 3.1 -- 3.3, 3.5; &nbsp; 8.1 -- 8.2, 8.4.1 -- 8.4.4

Several steps in a decryption algorithm need to see if a string is a valid English word.  We need a Python data structure that allows for a fast lookup in a collection of 100,000 or more strings.  

We could store the words in a list, but (as you should recall from CIS 210) the amount of time to find a string in a list is proportional to the length of the list.

In Chapter 8 there is a suggestion to use a dictionary instead of a list.  Each key is an English word and the value can be anything, *e.g.* the number 0:

``` 
words = { 'a' : 0, 'aardvark' : 0, 'abaci' : 0, ... }
```

To see if a string is an English word just check to see if the string is a key, using the `in` operator:
```
if w in words:
   ...
```
The time for this operation is proportional to $log_2 n$ where $n$ is the number of words in the dictionary.

What about sets?  Do you think Python organizes set objects so it can do a $log_2$ time lookup?  This project will help you find out.

| data structure | time to search |
| ----- | ----- |
| List | $\mathcal{O}(n)$ |
| Dictionary | $\mathcal{O}(\log n)$ |
| Set | $\mathcal{O}(\log n)$ ?|


wordlist -- return list of words from file
sample(wl, n) -- return dictionary {'list': lst, 'dict': dct, 'set': set} with randomly chosen words
test(x, m) -- pass sample, compute time to look up random collection of words (one is too short)
extra -- show how to plot results of experiment (saved in data.csv)

###  <span style="color:teal">Part 1: &nbsp; Word List (N points)</span>

A file named `wordlist.txt` (you'll find a link in the project description on Canvas) has a list of English words, with one word per line.

Use the code cell below to write Python code that opens the file, reads the words, and saves them in a list named `wordlist`.

##### <span style="color:red">Code</span>

<!-- grade_id: wordlist, points: 5.0, grade: true, solution: true -->
```python
# YOUR CODE HERE
pass
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: wordlist_test_1, points: 5.0, grade: true, solution: false -->
```python
# there are 106,566 words in the file
assert len(wordlist) == 106566
```

<!-- grade_id: wordlist_test_2, points: 5.0, grade: true, solution: false -->
```python
# look up some words we expect to be in the dictionary
for w in ['python', 'programming', 'in', 'context']:
    assert w in wordlist
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: wordlist_doc, points: 5.0, grade: true, solution: true -->
YOUR DOCUMENTATION HERE

###  <span style="color:teal">Part 2: &nbsp; Lists, Dictionaries, Sets (N points)</span>

Write a function named `wordlist_sample` that has two arguments, a list of words and an integer `n`.  The function should select `n` words at random from the word list created by the previous project and then make a list, a dictionary, and a set that all contain the same words.  The return value will be a tuple containing the three data structures.

**Note:** to select words at random from a list use the `sample` function from Python's `random` library.

Example:
```
>>> wl, wd, ws = wordlist_sample(5)

>>> print(wl); print(wd); print(ws)
['chronometry', 'vicariousness', 'mischarged', 'gadding', 'languishing']
{'languishing': 0, 'vicariousness': 0, 'gadding': 0, 'chronometry': 0, 'mischarged': 0}
{'gadding', 'languishing', 'vicariousness', 'chronometry', 'mischarged'}
```


##### <span style="color:red">Code</span>

<!-- grade_id: wordlist_sample, points: 8, grade: true, solution: true -->
```python
# YOUR CODE HERE
pass
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: sample_test_1, points: 4, grade: true, solution: false -->
```python
# Make a list of 10 words, verify the items passed back have the expected types

wl, wd, ws = wordlist_sample(10)
assert isinstance(wl, list)
assert isinstance(wd, dict)
assert isinstance(ws, set)
```

<!-- grade_id: sample_test_2, points: 4, grade: true, solution: false -->
```python
# Check to see if the list and dictionary have the same contents

wl, wd, ws = wordlist_sample(10)
assert len(wl) == len(wd)
for x in wl:
    assert x in wd
```

<!-- grade_id: sample_test_3, points: 4, grade: true, solution: false -->
```python
# Same as above, but compare the list and the set

wl, wd, ws = wordlist_sample(10)
assert len(ws) == len(ws)
for x in wl:
    assert x in ws
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: wordlist_xxx, points: 5, grade: true, solution: true -->
YOUR DOCUMENTATION HERE

###  <span style="color:teal">Part 3: Experiments</span>

This section of the notebooks shows off some of the features of Jupyter and introduces `matplotlib`, a scientific visualization library we'll use later in the term.

I ran some experiments to measure how long it takes to look up words in each data structure.  If we simply look up a random word $w$ it's possible the $w$ could be the very first word in the data structure and there would be no difference.  Similarly, if the random sample has a lot of words that start with 'a' and the list is sorted the list lookup times will be artificially short.  So in these experiments I did searches for a set of "nonsense words" -- words I know are not in the word list, and which guarantee the search takes the longest time for each type of data structure.

The code cell below defines a special type of class called a "context manager" and a helper function named `time_lookups` that uses the context manager to measure execution time.  Pass `time_lookups` a list of words to look for and either a list, dictionary, or set and it will return the execution time (in milliseconds) it took to look for all the words in that data structure.

```python
import time

class Timer:
    '''
    A Timer object is a context manager that can be used to measure execution times.
    '''
    def __enter__(self):
        self._start = time.time()
        return self
    
    def __exit__(self, *args):
        self._stop = time.time()
        self.msecs = 1000 * (self._stop - self._start)
        
def time_lookups(lst, struct):
    '''
    Return the amount of time it takes to look up all the strings in lst
    in a container (a set, dictionary, or list).  The return value is the
    average execution time measured over n tests.
    '''
    num_found = 0
    with Timer() as t:
        for w in lst:
            if w in struct:     # note the 'in' operator works for all 3 data structures
                num_found += 1
    assert num_found == 0       # make sure each lookup fails
    return t.msecs
```

The code cell below illustrates how to use `time_lookups` to measure how long it takes to search for the nonsense words in a random list of 100 words.  Note that to make a fair comparison we want each data structure to have the same words (which is why we wrote the `wordlist_sample` function).

```python
wl, wd, ws = wordlist_sample(1000)
nonsense = ['bryllyg', 'slythy', 'toves', 'mimsy', 'borogroves']
print('list:', time_lookups(nonsense, wl))
print('dict:', time_lookups(nonsense, wd))
print('set: ', time_lookups(nonsense, ws))
```

This function runs an experiment by making several different sized data structures and recording how long it takes to search for nonsense words in each structure.  The default experiment is based on 10 different sizes from 10 up to 100.  

For each size the function calls `wordlist_size` to make three containers of that size and then save the execution time for looking up words in each container.  The results are then plotted and displayed in the notebook.

```python
import matplotlib.pyplot as plt
% matplotlib inline

def run_experiment(minsize = 10, maxsize = 100, stepsize = 10, scale = 1):
    
    tl = []     # time to search a list
    td = []     # time to search a dictionary
    ts = []     # time to search a set
    x = []      # structure size (x-axis in plot)
    
    nonsense = ['jabberwock', 'bryllyg', 'slythy', 'toves', 'mimsy', 'borogroves', 
                'outgrabe', 'vorpal', 'frabjous', 'jubjub', 'raths', 'uffish',
                'tulgey', 'frumious', 'manxome', ]
    
    for size in range(minsize, maxsize+1, stepsize):
        wl, wd, ws = wordlist_sample(size)
        tl.append(time_lookups(nonsense, wl))
        td.append(scale * time_lookups(nonsense, wd))
        ts.append(scale * time_lookups(nonsense, ws))
        x.append(size)
        
    plt.plot(x, tl)
    plt.plot(x, td)
    plt.plot(x, ts)
    plt.legend(('list', 'dict', 'set'), loc='upper left')
    plt.axis([0, maxsize+minsize, 0, max(tl)*1.1])
    plt.xlabel('number of words')
    plt.ylabel('time (msec)')

```

```python
run_experiment()
```

This experiment tests 50 different sizes ranging from 1,000 up to 50,000 (about half the words in the dictionary).  

```python
run_experiment(minsize=1000, maxsize=50000, stepsize=1000)
```

Even with 50,000 words it hardly takes any time at all to look for words in dictionaries and sets.  Pass a "scale" parameter to the function to tell it to multiply the times for those two structures by a constant factor so we can "zoom in".

```python
run_experiment(minsize=1000, maxsize=50000, stepsize=1000, scale = 100)
```

### Extra Credit Challenges

(1) Do the results from the last plot seem accurate?  Can it really be the case that looking up words in a dictionary or set with 50,000 items is roughly as fast as looking up the same words in a dictionary or set with 1,000 items? Write an explanation for why the dictionary and set times are this low. Hint: do some "back of the envelope" calculations.  What is $\log_2 50000$?

(2) Why are the lines in the plots so "messy"?  Why do the lines go down some times, instead of always increasing from one size to the next?

(3) Modify the `run_experiment` function so it returns the average time from several measurements.  Add an extra parameter named `repetitions` with a default of 1.  Then for each problem size run the specified number of repetitions and save the mean execution time in the result lists.

Mail your explanations/notebooks to **cis211-extra@cs.uoregon.edu**.



```python

```