# <span style="color:teal;">CIS 211 Project 2:  Cryptography</span>

##### Due 11:00 P.M. April 14, 2017

##### Reading:  M&R 3.1 -- 3.3, 3.5; &nbsp; 8.1 -- 8.2, 8.4.1 -- 8.4.4

This week we're going to use simple encryption and decryption algorithms as a way of exploring strings and containers (lists, dictionaries, and sets) in Python.

Encryption is described in Chapter 3 of the textbook.  You should understand how the **substitution cipher** works -- you don't need to understand the details of any of the Python functions that implement the method, but you should know what it means to encode a piece of text with this kind of cipher.

Writing a function that decrypts a piece of text (assuming it really was encrypted with a substitution cipher) is not too hard, compared to other forms of encryption, but still more than we want to do in a one-week project.  Instead of developing a complete algorithm that will decrypt a piece of text we'll write a few small "helper functions" that might be useful as part of a larger application.


###  <span style="color:teal">Imports</span>

```python
from string import punctuation, ascii_uppercase, ascii_lowercase
```

###  <span style="color:teal">Test Data</span>

```python
seuss = '\nUEK XZIV, MOU XZIV,\nRKP XZIV, QHGK XZIV.\n'
print(seuss)
```

Output:
```
UEK XZIV, MOU XZIV,
RKP XZIV, QHGK XZIV.
```

###  <span style="color:teal">Project 1: &nbsp; Ciphertext Words (20 points)</span>

Fill in the body of the function named `cipher_words`.  It will be passed a string containing a "cryptoquote", and it should return the set of words in the quote.  All punctuation should be stripped from the ends of words (but not in the middle).

Example:
```
>>> cipher_words(suess)
{'MOU', 'QHGK', 'RKP', 'UEK', 'XZIV'}
```

**Style (2 points):** This function can be written without a `for` loop, using set comprehension.

##### <span style="color:red">Code</span>

<!-- grade_id: cipher_words, points: 10, grade: true, solution: true -->
```python
def cipher_words(text):
    ''' (str) -> dict
    The function takes a string as a parameter and then converts the string into a set of words without any punctuation.
    '''
    words = {w.strip(punctuation) for w in text.split()}
    return words
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: cipher_words_test, points: 5, grade: true, solution: false -->
```python
# pass the example text to the function, make sure 5 words are returned
words = cipher_words(seuss)
assert len(words) == 5

# look up each word in the result set
for w in ['MOU', 'QHGK', 'RKP', 'UEK', 'XZIV']:
    assert w in words
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: cipher_words_doc, points: 5, grade: true, solution: true -->
The function saves the set comprehension to words. Within the comprehension, the text argument is split on the whitespace and then each word in words is stripped of its punctuation. The set of words is then returned.

###  <span style="color:teal">Project 2: &nbsp; Count Ciphertext Letters (20 points)</span>

Fill in the body of the `count_unsolved` function.  It will be passed a single word, and it should return the number of unique ciphertext (upper case) letters in the word.

Examples:
```
>>> count_unsolved('XZIV')
4

>>> count_unsolved('QTETET')
3
```
Note how there are 6 letters in the second example, but some are repeated.  There are only 3 different letters:  Q, E, and T.

##### <span style="color:red">Code</span>

<!-- grade_id: count_unsolved, points: 10, grade: true, solution: true -->
```python
def count_unsolved(word):
    '''(str) -> int
    The function takes a string as a parameter and returns the count of upper case characters within the string as an integer.
    '''
    count = []
    for ch in word:
        if ch not in count and ch.isupper():
            count.append(ch)
            
    return len(count)
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: count_unsolved_test_1, points: 1, grade: true, solution: false -->
```python
# each letter in this string is unique
assert count_unsolved('XZIV') == 4
```

<!-- grade_id: count_unsolved_test_2, points: 1, grade: true, solution: false -->
```python
# there are only 3 unique letters in this string
assert count_unsolved('QTETET') == 3
```

<!-- grade_id: count_unsolved_test_3, points: 1, grade: true, solution: false -->
```python
# repeat the previous test after "solving" one of the letters
assert count_unsolved('QaEaEa') == 2
```

<!-- grade_id: count_unsolved_test_4, points: 1, grade: true, solution: false -->
```python
# a word can be completely solved (no remaining ciphertext letters)
assert count_unsolved('banana') == 0
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: count_unsolved_doc, points: 6, grade: true, solution: true -->
The count is a list of the unsolved uppercase characters within the word. A for loop cycles through the word and appends the characters that are uppercase and are not already in the list count. After the for loop is done the length of unsolved characters is returned.

###  <span style="color:teal">Project 3: &nbsp; Fewest Unsolved Letters (15 points)</span>

One strategy for choosing words to work on is to find the word that has the fewest letters remaining to be solved.  Fill in the body of the `fewest_unsolved` function so it iterates over a set of words and returns a tuple `(n,w)` where `w` is the word with the fewest remaining undecrypted letters and `n` is the number of undecrypted letters.

For example, suppose the quote from earlier has been partially decrypted and the set of words is now
```
{'MOo', 'QHue', 'ReP', 'XiIV', 'oEe'}
```
Passing this set to `fewest_unsolved` should return `oEe` because it has one uppercase letter and the others all have two:
```
```

**Note:** If there is a tie (more than one word has the same minimum number of undecrypted letters) just return one at random (_e.g._ the first one found).

##### <span style="color:red">Code</span>

<!-- grade_id: fewest_unsolved, points: 6, grade: true, solution: true -->
```python
def fewest_unsolved(group):
    '''(set) -> str
    The function takes a set of words as an argument and returns the word within the set that has the fewest unsolved
    characters as a string. It uses the helper function count_unsolved().
    '''
    count = len(max(group))
    few = None
    for word in group:
        if count_unsolved(word) < count:
            count = count_unsolved(word)
            few = word
            
    return few

```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: fewest_unsolved_test1, points: 2, grade: true, solution: false -->
```python
# this set has a word with one ciphertext letter and three others with two ciphertext letters
test_set = {'oEe', 'XiIV', 'MOo', 'QHue', 'ReP'}

# the first call should find the word with one ciphertext letter
w = fewest_unsolved(test_set)
assert w == 'oEe'
```

<!-- grade_id: fewest_unsolved_test_2, points: 2, grade: true, solution: false -->
```python
# remove that word from the set, the next call should find a word with two ciphertext letters
test_set.remove(w)
w = fewest_unsolved(test_set)
assert count_unsolved(w) == 2
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: fewest_unsolved_doc, points: 5, grade: true, solution: true -->
First the count is intialized at the maximum length of the words in the group given and the string few is initialized as None. A for loop then iterates through the group of words and uses the helper function count_unsolved() to determine if the number of unsolved characters within the word is less than the count. If this is true then the word is saved as a string to few. And so on and so forth... The function then returns few as a string which is the word with the fewest unsolved characters.

###  <span style="color:teal">Project 4: &nbsp; Letter Pairs (30 points)</span>

An effective strategy for decrypting English text is to consider letters two at a time (see Section 8.4.3 of the textbook for an explanation of why this approach is helpful).

For this part of the project, fill in the definition of the function named `make_letter_pairs`.  It will be passed a string, and it should return a matrix $m$, such that $m_{i,j}$ is the frequency with which letters $i$ and $j$ were next to each other in the string.

One way to create a matrix in Python is to use a list of lists, but for this project we'll make a dictionary of dictionaries.  If we save the result of the function call in a variable named `m`, we can look up a frequency using indexing.  For example, we can look up the frequency of `'i'` followed by `'e'` by evaluating
```
>>> m['i']['e']
```

Here is an example:
```
>>> m = make_letter_pairs('how now, brown cow?')
>>> m
{'b': {'r': 0.1},
 'c': {'o': 0.1},
 'h': {'o': 0.1},
 'n': {'o': 0.1},
 'o': {'w': 0.4},
 'r': {'o': 0.1},
 'w': {'n': 0.1}}
```

The 10 places in this text where there are two letters next to each other are `ho`, `ow`, `no`, `ow`, `br`, `ro`, `ow`, `wn`, `co`, and `ow`.  The combination of `b` followed by `r` occurs only once, so its frequency is 1 / 10 = 0.1:
```
>>> m['b']['r']
0.1
```

The pair `ow` occurs four times, so its frequency is 4 / 10 = 0.4:
```
>>> m['o']['w']
0.4
```

**Note**  There is a similar function, named `neighborCount`, defined in the book. You can use this function to get some ideas for how to write your code, but there are some important differences you should be aware of:
* the function in the book creates counts, but our function creates frequencies
* the function in the book counts in both directions, _e.g._ when it sees the pair `ow` it adds 1 to the count of `ow` and one to the count for `wo`; we count one direction only

##### <span style="color:red">Code</span>

<!-- grade_id: letter_pairs, points: 15, grade: true, solution: true -->
```python
def make_letter_pairs(text):
    ''' (str) -> dict
    This function takes a string of text as an argument and then strips, lowercases, and removes all the punctuation 
    from the text and then splits the text into words. The function preforms a list comprehension on the words to create
    a list of the pairs. Then the dictionaries are formed together with the pairs and the frequencies to return a 
    dictionary of dictionaries of the pairs and their frequency in the orignal text. 

    '''
    words = [ch.strip(punctuation).lower() for ch in text.split()]

    lst_of_pairs = [ch[i] + ch[i+1] for ch in words for i in range(len(ch)-1) if ch.isalpha()]
    length_of_pairs = len(lst_of_pairs)
    freq_dict = {pair: (lst_of_pairs.count(pair) / length_of_pairs) for pair in lst_of_pairs}
    
    pairs = {ch[0]: {ch[-1]: freq_dict[ch]} for ch in freq_dict.keys()}
    for ch in pairs:
        pairs.setdefault(ch,{})    
    return pairs
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: letter_pairs_test_1, points: 4, grade: true, solution: false -->
```python
# the first test uses the 'brown cow' example
m = make_letter_pairs('how now, brown cow?')

# there are 7 different letters in the string
assert len(m) == 7

# check two of the frequencies
assert round(m['o']['w'],1) == 0.4
assert round(m['r']['o'],1) == 0.1
```

<!-- grade_id: letter_pairs_test_2, points: 2, grade: true, solution: false -->
```python
# the remaining tests are based on frequencies using a large body of text (War of the Worlds, by H. G. Wells)
with open('wells.txt') as f:
    pp = make_letter_pairs(f.read())

# look up some common letter combinations ("pp" stands for "pair probability")
assert round(pp['e']['e'], 3) == 0.006
assert round(pp['s']['s'], 3) == 0.004
```

Output:
```
<ipython-input-13-e040d921ee28> in make_letter_pairs(text)
     11     lst_of_pairs = [ch[i] + ch[i+1] for ch in words for i in range(len(ch)-1) if ch.isalpha()]
     12     length_of_pairs = len(lst_of_pairs)
---> 13     freq_dict = {pair: (lst_of_pairs.count(pair) / length_of_pairs) for pair in lst_of_pairs}
     14 
     15     pairs = {ch[0]: {ch[-1]: freq_dict[ch]} for ch in freq_dict.keys()}

<ipython-input-13-e040d921ee28> in <dictcomp>(.0)
     11     lst_of_pairs = [ch[i] + ch[i+1] for ch in words for i in range(len(ch)-1) if ch.isalpha()]
     12     length_of_pairs = len(lst_of_pairs)
---> 13     freq_dict = {pair: (lst_of_pairs.count(pair) / length_of_pairs) for pair in lst_of_pairs}
     14 
     15     pairs = {ch[0]: {ch[-1]: freq_dict[ch]} for ch in freq_dict.keys()}

KeyboardInterrupt:
```

<!-- grade_id: letter_pairs_test_3, points: 2, grade: true, solution: false -->
```python
# look up some rare combinations (in this book, at least)
assert round(pp['r']['z']) < 0.000001
assert round(pp['y']['n']) < 0.000001
```

<!-- grade_id: letter_pairs_test_4, points: 2, grade: true, solution: false -->
```python
# letter combinations 'jx' and 'kk' never occur in this book
assert 'x' not in pp['j']
assert 'k' not in pp['k']
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: letter_pairs_doc, points: 5, grade: true, solution: true -->
```python
'''
    This function takes a string of text as an argument and then strips, lowercases, and removes all the punctuation 
    from the text and then splits the text into words. The function preforms a list comprehension on the words to create
    a list of the pairs. Then the dictionaries are formed together with the pairs and the frequencies to return a 
    dictionary of dictionaries of the pairs and their frequency in the orignal text. 
'''
```

###  <span style="color:teal">Project 5: `apply` (15 points)</span>

The final part of the project is to write a function named `apply` that will use a decryption key (a dictionary mapping ciphertext letters to plaintext) to update a partially decrypted string.

For example, this key replaces X with `'f'`, Z with `'i'`, I with `'s'`, and V with `'h'`:
```
>>> k = {'X':'f', 'Z':'i', 'I':'s', 'V':'h'}
```

If we apply the key to the example ciphertext this is the result:
```
>>> apply(k, 'UEK XZIV, MOU XZIV, RKP XZIV, QHGK XZIV.')
'UEK fish, MOU fish, RKP fish, QHGK fish.'
```

##### <span style="color:red">Code</span>

<!-- grade_id: apply, points: 6, grade: true, solution: true -->
```python
def apply(key, quote):
    text = [ch.strip(punctuation) for ch in quote.split()]
    print(text)
    lett_dict = {ch for ch in text}
    print(lett_dict)
    pass
apply(k,'UEK XZIV, MOU XZIV, RKP XZIV, QHGK XZIV.')
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: apply_test, points: 4, grade: true, solution: false -->
```python
# make a decryption key to use in the tests
k = {'X':'f', 'Z':'i', 'I':'s', 'V':'h'}
apply(k,'UEK XZIV, MOU XZIV, RKP XZIV, QHGK XZIV.')
# apply the key 
assert apply(k,'XZIV') == 'fish'
assert apply(k,'UDKLCXZIV') == 'UDKLCfish'
assert apply(k,'XMIVZQP') == 'fMshiQP'
assert apply(k,'JSBWC') == 'JSBWC'
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: apply_doc, points: 5, grade: true, solution: true -->
YOUR DOCUMENTATION HERE