# <span style="color:teal;">CIS 211 Project 5:  Building GUIs with Tk</span>

##### Due 11:00 P.M. Friday May 5

Reading:
* [TkDocs - Tk Tutorial](http://www.tkdocs.com/tutorial/index.html)
* [Tkinter Tutorial](http://www.python-course.eu/python_tkinter.php) from `python-course.eu`

### <span style="color:teal">Setup</span>

Execute the following code cell to tell IPython you want to create windows with Tk and to import the `tkinter` library.

```python
%gui tk
import tkinter as tk
```

<span style="color:red">**Important:**</span> &nbsp; The autograder tests assume `tkinter` has been imported as shown above.  You should follow the same convention, and refer to widget classes as `tk.Label`, `tk.Entry`, and so on.

<span style="color:red">**Important:**</span> &nbsp; The code cells where you define your GUIs contain assignment statements that create global variables.  For example, the first GUI is assigned to a variable named `hello_app`.  **Do not change these assignment statements**.  The autograder test cells will use these variables to access parts of your GUIs.

### <span style="color:teal">1. &nbsp; Hello, World (10 points)</span>

The code cell for this project has an assignment statement that creates a top level Tk application named `hello_app`.  

Add statements to the code cell so the program has a label that says "Hello, World":

<img src="http://www.cs.uoregon.edu/Classes/15S/cis211/images/hello_gui.png"/>

The example above uses 24-point Helvetica.  You can use the default text, but feel free to experiment with other fonts, font sizes, colors, *etc*.

The goal for this part of the project is just to make sure you have all the necessary software installed and can create a simple application.  There is no documentation requirement for this part.

##### <span style="color:red">Code:</span>

<!-- grade_id: hello, points: 4, grade: true, solution: true -->
```python
hello_app = tk.Tk()
label = tk.Label(hello_app, padx=20, pady=10, text='Hello, World!', font='Times 24 bold')
label.pack()
```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: hello_test_1, points: 2, grade: true, solution: false -->
```python
# The GUI should have a single widget
widgets = list(hello_app.children.values())
assert len(widgets) == 1
```

<!-- grade_id: hello_test_2, points: 2, grade: true, solution: false -->
```python
# Check to see if the widget is a Label
widgets = list(hello_app.children.values())
assert isinstance(widgets[0], tk.Label)
```

<!-- grade_id: hello_test_3, points: 2, grade: true, solution: false -->
```python
# The text in the label should start with "hello" (upper or lower case)
widgets = list(hello_app.children.values())
assert widgets[0]['text'].lower().startswith('hello')
```

### <span style="color:teal">2. &nbsp; Payment Calculator (50 points)</span>

Create a program that computes the monthly payment on a loan. The payment is a function of the loan amount, the annual interest rate, and the number of years to pay off the loan. To calculate the payment, first compute two values named $r$ and $p$, defined by these equations:

$ r = ( \textrm{rate} \; / \; 100 ) \; / \; 12 $

$ p = 12 \times \textrm{years} $

The monthly payment is then

$$ \textrm{payment} = \frac{ r \times \textrm{amount} }{ 1 - (1 + r)^{-p} } $$

For example, if you take out a 30-year loan for \$150,000 at an annual rate of 4.5% your payment would be \$760.03 per month.

The GUI for this application should have four text entry boxes, and each box should have a label to its left that explains what the box holds.  The first three boxes are for the user to enter the loan amount, the interest rate, and the number of years to repay the loan. 

The GUI should have a button labeled “compute payment”. When the user clicks this button, your program should read the values of the three parameters, use the equations shown above to calculate the monthly payment, and display the payment amount in the fourth text entry box.

<img src="http://www.cs.uoregon.edu/Classes/15S/cis211/images/payment_gui.png"/>

The code cell for this program has the outline for the program.  First fill in the body of the function named `monthly_payment` and make sure it computes the correct payment value given an amount, rate, and loan period.  When your write the button callback function have it call the `monthly_payment` function.

Next add statements that create the labels, boxes, and button for the application.  When you create the Entry and Button widgets, make sure you save them in global variables with the following names so the autograder tests will succeed:
* `amount`
* `rate`
* `years`
* `payment`
* `button`

##### <span style="color:red">Code:</span>

<!-- grade_id: payment, points: 15, grade: true, solution: true -->
```python
def monthly_payment(amt, pct, yrs):
    '''(float, float, int) -> float
    
    '''
    r = (pct/100)/12
    p = 12 * yrs
    pymt = round((r * amt)/(1 - pow((1 + r),(-p))), 2)
    return pymt
    

def compute_payment_cb():
    payment.delete(0, tk.END)
    if (len(amount.get()) > 0) and (len(rate.get()) > 0) and (len(years.get()) > 0):
        payment.delete(0, tk.END)
        pymt = monthly_payment(float(amount.get()), float(rate.get()),int(years.get()))
        payment.insert(0, '${:.2f}'.format(pymt))
        return
        
    amount.delete(0, tk.END)
    rate.delete(0, tk.END)
    years.delete(0, tk.END)
    return
    
    

payment_app = tk.Tk()

amount_label = tk.Label(payment_app, text='Loan Amount:',font='Times 12')
amount_label.grid(row=0, column=0, sticky='W', padx=30, pady=10)

amount = tk.Entry(payment_app, font='Times 12')
amount.grid(row=0, column=1, sticky='E', padx=30, pady=10)

rate_label = tk.Label(payment_app, text='Interest Rate:',font='Times 12')
rate_label.grid(row=1, column=0, sticky='W', padx=30, pady=10)

rate = tk.Entry(payment_app, font='Times 12')
rate.grid(row=1, column=1, sticky='E', padx=30, pady=10)

years_label = tk.Label(payment_app, text='Loan Periods:',font='Times 12')
years_label.grid(row=2, column=0, sticky='W', padx=30, pady=10)

years = tk.Entry(payment_app, font='Times 12')
years.grid(row=2, column=1, sticky='E', padx=30, pady=10)

payment_label = tk.Label(payment_app, text='Payment:',font='Times 12')
payment_label.grid(row=3, column=0, sticky='W', padx=30, pady=10)

payment = tk.Entry(payment_app, font='Times 12')
payment.grid(row=3, column=1, sticky='E', padx=30, pady=10)

button = tk.Button(payment_app, text='Compute Payment', font='Times 12', justify='center', command=compute_payment_cb)
button.grid(row=4, column=0, columnspan=2, pady=30, sticky='S')
```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: payment_test_1, points: 5, grade: true, solution: false -->
```python
# Test the callback function to make sure it is computing payments accurately
assert round(monthly_payment(20000, 3.0, 6), 2) == 303.87
```

<!-- grade_id: payment_test_2, points: 10, grade: true, solution: false -->
```python
# Verify the number of widgets of each type
counts = { tk.Label: 0, tk.Entry: 0, tk.Button: 0 }

for x in payment_app.children.values():
    counts[type(x)] += 1
    
assert counts == { tk.Label: 4, tk.Entry: 4, tk.Button: 1 }
```

<!-- grade_id: payment_test_3, points: 10, grade: true, solution: false -->
```python
# Store values in the GUI, click the button, check the result
amount.insert(0, '20000')
rate.insert(0, '3.0')
years.insert(0, '6')
button.invoke()

assert payment.get() == '$303.87'
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: doc_2, points: 10, grade: true, solution: true -->
First I converted the equation for monthly payment into code in the function monthly_payment() by calling the pow() function. Then I created the compute_payment_cb function, that clears the entry box for payment. Then if there are variables in all three of the entry boxes amount, rate, years, then the monthly_payment function is called to calculate and the total. The int total is then formatted with a dollar sign and as a float with 2 decimal places is inserted into the payment box. If any of the three parameters is left empty by the user the function clears all the boxes when the calculate button is clicked.
I then created label widgets and entry widgets for the amount, rate, years, and payment. As well as centered a button widget at the bottom of the GUI. The widgets are set up on a grid with 0-4 rows and 0-2 columns. All of the entry widgets are in column 2 and all the label widgets are in column 0, leaving the button widget for column 1. 

### <span style="color:teal">3. &nbsp; Canvas (40 points)</span>

Fill in the application named `canvas_app`.  The application will have just one widget, a Tk canvas with a size of $400 \times 400$ pixels. 

You also need to fill in the body of the function named `make_circle` that will draw a circle on your canvas.  The parameters are a reference to the canvas widget, three integers for the size and location of the circle, and a string with a color name.

Use the function to draw three circles on your canvas.  You can choose whatever size, location, and color you would like.

<img src="http://www.cs.uoregon.edu/Classes/16W/cis211/images/gui/three_circles.png"/>


##### <span style="color:red">Code:</span>

<!-- grade_id: canvas, points: 15, grade: true, solution: true -->
```python
def make_circle(canvas, radius, x, y, color):
    '''(reference_to_widget,int,int,int,str) -> ()
    The function uses the reference to a canvas widget and creates a circle by calling the create_oval method and
    by giving the other four arguments, radius, x, y, and color to that method as the rest of its parameters.
    The circle is drawn on a canvas but nothing is returned
    '''
    canvas.create_oval(x-radius,y-radius,x+radius,y+radius, fill=color)
    return

canvas_app = tk.Tk()
cvs = tk.Canvas(canvas_app, width=400, height=400)
cvs.pack()
make_circle(cvs, 50, 100, 100, 'green')
make_circle(cvs, 25, 200, 200, 'blue')
make_circle(cvs, 15, 300, 300, 'yellow')
```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: canvas_test_1, points: 5, grade: true, solution: false -->
```python
# See if the GUI has a Canvas widget
widgets = list(canvas_app.children.values())
assert len(widgets) == 1 and isinstance(widgets[0], tk.Canvas)
```

<!-- grade_id: canvas_test_2, points: 10, grade: true, solution: false -->
```python
# See if the canvas has circles on it
widgets = list(canvas_app.children.values())
for n in widgets[0].find_all():
    assert widgets[0].type(n) == 'oval'
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: doc_3, points: 10, grade: true, solution: true -->
First I defined the make_circle function so that it called the create_oval() method of our canvas widget that has been given as one of the parameters of this function. The method create_oval() uses the radius, x and y coordinates, and the color to draw a circle on our canvas widget that has been assigned and packed globally.
The last step was to call our new function three times with different x, y, and color parameters each time.